Heterotypic interfacial tension between oncogenic and wild-type populations forms the mechanical basis of tissue-specific oncogenesis in epithelia

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The behaviour of cells expressing constitutively active HRas is examined in mosaic monolayers, both in MCF10a breast epithelial and Beas2b bronchial epithelial cell lines, mimicking the potential initial phase of development of carcinoma. Single HRas-positive cells are excluded from MCF10a but not Beas2b monolayers. Most interestingly, however, when in groups, these cells are not excluded, but rather sharply segregated within a MCF10a monolayer. In contrast, they freely mix with wt Beas2b cells. Biophysical analysis identifies high tension at heterotypic interfaces between HRas and wild-type cells as the likely reason for segregation of MCF10a cells. The hypothesis is supported experimentally, as myosin inhibition abolishes segregation. The probable reason for lack of segregation in the bronchial epithelium is to be found in the different intrinsic properties of these cells, which form a looser tissue with lower basal actomyosin activity. The behaviour of single cells and groups is recapitulated in a vortex model based on the principle of differential interfacial tension, under the condition of high heterotypic interfacial tension.

Strengths:

Despite being long recognized as a crucial event during cancer development, segregation of oncogenic cells has been a largely understudied question. This nice work addresses the mechanics of this phenomenon through a straightforward experimental design, applying the biophysical analytical approaches established in the field of morphogenesis. Comparison between two cell types provides some preliminary clues on the diversity of effects in various cancers.

Weaknesses:

Although not calling into question the main message of this study, there are a few issues that one may want to address:

(1) One may be careful in interpreting the comparison between MCF10a and Beas2b cells as used in this study. The conditions may not necessarily be representative of the actual properties of breast and bronchial epithelia. How much of the epithelial organization is reconstituted under these experimental conditions remains to be established. This is particularly obvious for bronchial cells, which would need quite specific culture conditions to build a proper bronchial layer. In this study, they seemed to be on the verge of a mesenchymal phenotype (large gaps, huge protrusions, cells growing on top of each other, as mentioned in the manuscript).

As an alternative to Beas2b, comparison of MCF10a with another cell line capable of more robust in vitro epithelial organization, but ideally with different adhesive and/or tensile properties, would be highly interesting, as it may narrow down the parameters involved in segregation of oncogenic cells.

(2) While the seminal description of tissue properties based on interfacial tensions (Brodland 2002) is clearly key to interpreting these data, the actual "Differential Interfacial Tension Hypothesis" poses that segregation results from global differences, i.e., juxtaposition of two tissues displaying different intrinsic tensions. On the contrary, the results of the present work support a different scenario, where what counts is the actual difference in tension ALONG the tissue boundary, in other words, that segregation is driven by high HETEROTYPIC interfacial tension. This is an important distinction that should be clarified.

(3) Related: The fact that actomyosin accumulates at the heterotypic interface is key here. It would be quite informative to better document the pattern of this accumulation, which is not clear enough from the images of the current manuscript: Are we talking about the actual interface between mutant and wt cells (membrane/cortex of heterotypic contacts)? Or is it more globally overactivated in the whole cell layer along the border? Some better images and some quantification would help.

(4) In the case of Beas2b cells, mutant cells show higher actin than wt cells, while actin is, on the contrary, lower in mutant MCF10a cells (Figure 2b). Has this been taken into account in the model? It may be in line with the idea that HRas may have a different action on the two cell types, a possibility that would certainly be worth considering and discussing.

Comments on revisions:

There is still one last point that should be made even clearer:

The system is being modelled based on the principle of INTERFACIAL TENSION, a description pioneered by the works of Steinberg and of Harris, and nicely conceptualized by Brodland (2002). Now the observed behaviour is a perfect case of sorting based on higher interfacial tension AT the boundary between cell types (with nice additional documentation of local actin and myosin enrichment in the revised manuscript). What needs to be made crystal clear it that this is NOT equivalent to the model of DITH ("DIFFERENTIAL INTERFACIAL TENSION HYPOTHESIS)" (Brodland 2002, Krieg et al 2008). It is important to stop using DITH in this context, as it leads to confusion and misinterpretations. Indeed, DITH predicts cell/tissue sorting based on differences in interfacial tension WITHIN the two cell types. While DITH accounts for relative POSITIONING (one tissue engulfing the other), it is now established that this is not the motor for cell sorting and tissue segregation, the key parameter is being heterotypic tension at the heterotypic interface. I thus invite the authors to avoid the terms "differential"/DITH, and rather use either "interfacial tension", or specifically to "HIGH HETEROTYPIC INTERFACIAL TENSION".

Related: the authors correctly cite Canty et al NatComm2017 when discussing this phenomenon. I suggest to add an additional key supporting reference "D.M. Sussman, J.M. Schwarz, M.C. Marchetti, M.L. Manning, Soft yet sharp interfaces in a vertex model of confluent tissue, Phys. Rev. Letters 120 (2018) 058001". One may also include another pioneer work in Drosophila is "M. Aliee, J.C. Roper, K.P. Landsberg, C. Pentzold, T.J. Widmann, F. Julicher, C. Dahmann, Physical mechanisms shaping the Drosophila dorsoventral compartment boundary, Curr. Biol. 22 (2012) 967-976."

We thank the reviewer for this important clarification. We fully agree that the mechanism underlying the observed segregation in our system is best described in terms of elevated heterotypic interfacial tension, rather than the classical Differential Interfacial Tension Hypothesis (DITH). As the reviewer correctly points out, DITH in its original formulation refers to differences in intrinsic interfacial tensions within each cell population, which primarily governs relative positioning (e.g., tissue engulfment), rather than the local sorting dynamics we observe here.

In contrast, our experimental and modeling results support a scenario in which segregation is driven by increased tension specifically at heterotypic interfaces between HRasV12 and wild-type cells. We agree that continued use of the term “Differential interfacial tension” in this context may lead to conceptual ambiguity.

Accordingly, we have revised the manuscript throughout to replace references to “differential interfacial tension” with more precise terminology, namely “interfacial tension” or “heterotypic interfacial tension”, wherever appropriate. We have also updated the Discussion to explicitly clarify this distinction and its implications for interpreting our results.

We thank the reviewer for suggesting additional relevant literature which have now included.

Reviewer #2 (Public review):

Summary:

The authors investigate the behavior of oncogenic cells in mammary and bronchial epithelia. They observe that individual oncogenic cells are preferentially excluded from the mammary epithelium, but they remain integrated in the bronchial epithelium. They also observe that clusters of oncogenic cells form a compact cluster in mammary epithelium, but they disperse in the bronchial epithelium. The authors demonstrate experimentally and in the vertex model simulations that the difference in observed behavior is due to the differential tension between the mutant and wild-type cells due to a differential expression of actin and myosin.

Strengths:

Very detailed analysis of experiments to systematically characterize and quantify differences between mammary and bronchial epithelia

Detailed comparison between the experiments and vertex model simulations to identify the differential cell line tension between the oncogenic and wild-type cells as one of the key parameters that are responsible for the different behavior of oncogenic cells in mammary and bronchial epithelia

Weaknesses:

It is unclear what is the mechanistic origin of the shape-tension coupling, which is used in the vertex model, and how important that coupling is for the presented results. Authors claim that the shape-tension coupling is due to the anisotropic distribution of stress fibers when cells are under external stress. It is unclear why the stress fibers should affect an effective line tension on the cell boundaries and why the stress fibers should be sensitive to the magnitude of the internal isotropic cell pressure. In experiments, it makes sense that stress fibers form when cells are stretched. Similar stress fibers form when cytoskeleton or polymer networks are stretched. It is unclear why the stress fibers should be sensitive to the magnitude of internal isotropic cell pressure. If all the surrounding cells have the same internal pressure, then the cell would not be significantly deformed due to that pressure and stress fibers would not form. Authors should better justify the use of the shape-tension coupling in the model, since most of the observed behavior is already captured by the differential tension even if there is no shape-tension coupling.

We thank the reviewer for this comment. We agree that we did not provide a mechanistic origin for the shape-tension coupling. In our model, stress fiber formation, along with actin ring formation, indicated that cells at the interface were elongated. Hence, we hypothesised that an interfacial force could induce nematic alignment at the interface. However, such an activity would only be feasible if the interface interaction were sufficiently high. Thus, the isotropic pressure at the heterotypic interface served as a proxy for cell-cell interactions in our model. However, inspired by recent work [1], we have tested whether activation of cells at the interface by shear stress would produce similar results. Exploring this aspect will require additional simulations.

(1) Pérez-Verdugo, F., Maniou, E., Galea, G. L., & Banerjee, S. (2026). Mechanosensitive feedback organizes cell shape and motion during hindbrain neuropore morphogenesis. Current Biology.

The observed difference of shape indices between the interfacial and bulk cells in simulations in the absence of differential line tension is concerning. This suggests that either there are not enough statistics from the simulations or that something is wrong with the simulations. For all presented simulation results, the authors should repeat multiple simulations and then present both averages and standard deviations. This way it would be easier to determine whether the observed differences in simulations are statistically significant.

The observed differences in shape indices between interfacial and bulk cells in simulations in the zero-line-tension case (Lambda=0) remain non-zero at the zero-stress threshold because the interface cells are still subject to the shape-dependent contribution gamma_ij, since the current model treats gamma_ij as independent of Lambda. We are exploring the possible relationship between Lambda and gamma_ij, and we will update this in the next version of the manuscript.

Recommendations for the authors:

The editor recommends considering the new comment made by reviewer #1 in his/her report:

"There is still one last point that should be made even more clear:

The system is being modelled based on the principle of INTERFACIAL TENSION, a description pioneered by the works of Steinberg and of Harris, and nicely conceptualized by Brodland (2002). Now the observed behaviour is a perfect case of sorting based on higher interfacial tension AT the boundary between cell types (with nice additional documentation of local actin and myosin enrichment in the revised manuscript). What needs to be made crystal clear it that this is NOT equivalent to the model of DITH ("DIFFERENTIAL INTERFACIAL TENSION HYPOTHESIS)" (Brodland 2002, Krieg et al 2008). It is important to stop using DITH in this context, as it leads to confusion and misinterpretations. Indeed, DITH predicts cell/tissue sorting based on differences in interfacial tension WITHIN the two cell types. While DITH accounts for relative POSITIONING (one tissue engulfing the other), it is now established that this is not the motor for cell sorting and tissue segregation, the key parameter is being heterotypic tension at the heterotypic interface. I thus invite the authors to avoid the terms "differential"/DITH, and rather use either "interfacial tension", or specifically to "HIGH HETEROTYPIC INTERFACIAL TENSION".

Related: the authors correctly cite Canty et al NatComm2017 when discussing this phenomenon. I suggest to add an additional key supporting reference "D.M. Sussman, J.M. Schwarz, M.C. Marchetti, M.L. Manning, Soft yet sharp interfaces in a vertex model of confluent tissue, Phys. Rev. Letters 120 (2018) 058001". One may also include another pioneer work in Drosophila is "M. Aliee, J.C. Roper, K.P. Landsberg, C. Pentzold, T.J. Widmann, F. Julicher, C. Dahmann, Physical mechanisms shaping the Drosophila dorsoventral compartment boundary, Curr. Biol. 22 (2012) 967-976."

Please see response to Reviewer 1

Reviewer #2 (Recommendations for the authors):

The authors have improved the manuscript and addressed some of my concerns. However, some of the questions were not adequately addressed.

(1) I appreciate additional justification regarding the need for the shape-tension coupling in the vertex model. However, the authors have not answered my question regarding why the shape-tension coupling model should be sensitive to the magnitude of the internal isotropic cell pressure. In experiments, it makes sense that stress fibers form when cells are stretched, but it is unclear why the stress fibers should be sensitive to the magnitude of internal isotropic cell pressure. If all the surrounding cells have the same internal pressure, then the cell would not be significantly deformed due to that pressure, and stress fibers would not form.

We thank the reviewer for pointing this out. We agree that we did not provide a mechanistic origin for the shape-tension coupling. In our model, stress fiber formation, along with actin ring formation, indicated that cells at the interface were elongated. Hence, we hypothesized that an interfacial force could induce nematic alignment at the interface. However, such an activity would only be feasible if the interface interaction were sufficiently high. Thus, the isotropic pressure at the heterotypic interface served as a proxy for cell-cell interactions in our model.

However, inspired by recent work [1], we have tested whether activation of cells at the interface by shear stress would produce similar results. Exploring this aspect will require additional simulations.

(1) Pérez-Verdugo, F., Maniou, E., Galea, G. L., & Banerjee, S. (2026). Mechanosensitive feedback organizes cell shape and motion during hindbrain neuropore morphogenesis. Current Biology.

(2) I appreciate that the authors provided additional statistics related to simulations. I am still very concerned about the observed difference in the shape indices between the cells at the interface and the bulk, when the interfacial line tension is exactly zero (Lambda=0). In that case, the cells at the interface and at the boundary are identical, and there should be no difference in the shape indices. Are cells at the interface for the zero-line tension case (Lambda=0) still subject to the shape dependent contribution gamma_ij? If that contribution is still included for the cells at the interface, then this could explain why cells at the interface are still different from cells in the bulk even when Lambda=0.

The observed differences in shape indices between interfacial and bulk cells in simulations in the zero-line-tension case (Lambda=0) remain non-zero at the zero-stress threshold because the interface cells are still subject to the shape-dependent contribution gamma_ij, since the current model treats gamma_ij as independent of Lambda. We are exploring the possible relationship between Lambda and gamma_ij, and we will update this in the next version of the manuscript.

(3) Authors included several additional supplemental figures (Figs. S4, S5, S6, S7) , but they are not discussed in the manuscript text. These new supplemental figures were only discussed in the rebuttal letter. These figures should also be discussed in the manuscript text.

We have cited the new supplementary figures in the main text.

(4) Authors have answered in the rebuttal letter what experimental data was used in Fig. 4c. This information also needs to be provided in the manuscript text.

We have added this information in the caption of Figure 4

(5) Supplementary Figure 3 is missing. That figure got moved to the appendix.

This has been rectified in the Supplementary file and the citations have been updated accordingly in the main text.

(6) At the end of section 4 in the main text, the authors introduced a new sentence regarding simulations of the vertex model with interfacial tension and mechanochemical feedback. The details of that model are described in the appendix, but it would be helpful to add a sentence or two already in the main text describing what is the mechanism of the mechanochemcial feedback.

We have added a line describing the mechanism of mechanochemical feedback.

(7) In the definition of the eccentricity, 'a' should be the minor axis and 'b' the major axis, i.e., 'a' and 'b' should be swapped.

We have corrected this.

(8) There is a typo at the end of the vertex model description in the methods section. "The details of the shape-tension coupling is described in the interface." The word interface should be an appendix.

We have fixed the typo.

(9) In the appendix section describing the shape-tension coupling, the authors should explain how the cell's director n is defined.

We have added a line in the appendix section describing shape-tension coupling explaining how the cell’s director n is defined.

(10) In Appendix Fig. 1, the two angles are defined as theta and theta' but the figure caption is defining angles theta_1 and theta_2. These angles need to be consistent.

This has been fixed.

Changes from the earlier version to the revised version

Abstract

Earlier version: Why does the same oncogenic mutation drive tumor formation in some tissues but not in others? While cancer driver mutations are well-documented, their tissue-specific effects remain largely attributed to genetic factors, leaving the biophysical aspects underexplored. Here, we demonstrate that mechanical interactions, specifically interfacial tension between newly transformed and wildtype epithelial cells are critical in determining survival and growth of HRasV12 oncogenic mutants in human mammary and bronchial epithelia, leading to contrasting outcomes in the two tissues. In mammary epithelium, isolated oncogenic cells are extruded- a typical mechanism of defense against cancer in epithelia- while oncogenic groups become spatially confined in a kinetically arrested, jammed state, marked by an actomyosin belt at the interface. In contrast, bronchial epithelium permits persistent spreading of the same oncogenic cells, which form long protrusions regardless of colony size. Furthermore, oncogenic clusters in these two tissues exhibit distinct biophysical properties, including variations in cell shapes, intracellular pressure, cell-cell tension, and cellular motility. Using a cell shape-tension coupled bi-disperse vertex model, we reveal that differences in interfacial tension at mutant–wild-type boundaries dictate whether oncogenic cells are eliminated, restrained, or expanded and that modulating interfacial tension alters mutant cell fate within the epithelium. Together, our findings uncover a mechanical basis for tissue-specific oncogenesis by highlighting how differential cellular mechanics at the oncogenic– host cell interface regulate tumor initiation and progression.

Revised version: Why does the same oncogenic mutation drive tumor formation in some tissues but not in others? While cancer driver mutations are well-documented, their tissue-specific effects remain largely attributed to genetic factors, leaving the biophysical aspects underexplored. Here, we demonstrate that mechanical interactions, specifically interfacial tension between newly transformed and wildtype epithelial cells are critical in determining survival and growth of HRasV12 oncogenic mutants in human mammary and bronchial epithelia, leading to contrasting outcomes in the two tissues. In mammary epithelium, isolated oncogenic cells are extruded-a typical mechanism of defense against cancer in epithelia-while oncogenic groups become spatially confined in a kinetically arrested, jammed state, marked by an actomyosin belt at the interface. In contrast, bronchial epithelium permits persistent spreading of the same oncogenic cells, which form long protrusions regardless of colony size. Furthermore, oncogenic clusters in these two tissues exhibit distinct biophysical properties, including variations in cell shapes, intracellular pressure, cell-cell tension, and cellular motility. Using a cell shape-tension coupled bi-disperse vertex model, we reveal that differences in interfacial tension at mutant–wild-type boundaries dictate whether oncogenic cells are eliminated, restrained, or expanded and that modulating the heterotypic interfacial tension alters mutant cell fate within the epithelium. Together, our findings uncover a mechanical basis for tissue-specific oncogenesis by highlighting how differential cellular mechanics at the oncogenic–host cell interface regulate tumor initiation and progression.

Introduction

Cancer has traditionally been understood to be a consequence of deregulated cell proliferation driven by genetic alterations that enable clonal expansion and tumor formation [1]. Interestingly, genomic analysis reveals that many tumor suppressor genes and oncogenes exhibit tissue specificity, meaning they are altered in some cancers but not others [2-4]. Current research suggests that a combination of intrinsic biological factors [4], such as differential epigenetic changes [3,5] and aneuploidy patterns [6-8], as well as extrinsic factors like the tumour microenvironment, including the presence of tissue-specific immune cells [9-10] influences tissue-specific cancer outcomes. However, a relatively underexplored yet crucial aspect of tissue-specific oncogenesis is the mechanical interaction between normal and cancerous cells. To this end, recent studies have highlighted a tumor suppressive behavior in epithelial tissues known as epithelial defense against cancer (EDAC) [11], where transformed cells are expelled from the tissue by surrounding normal cells [12-14]. However, EDAC can fail [14], allowing mutant cells to evade extrusion when the microenvironment becomes tumour-permissive, such as in conditions of increased stiffness [15] or inflammation due to a high-fat diet [16]. The success or failure of EDAC is closely linked to changes in cellular contractility [12-14] and cell-cell adhesions [17], suggesting that the underlying biophysical properties of the tissue play a decisive role in determining the outcome of cell competition. Since these properties can vary across different epithelial tissues, tissue-specific differences in cancer susceptibility may be explained by the varying ability of epithelia to manage newly transformed cells. However, despite these advances, the mechanical basis of tissue-specific oncogenesis remains missing. Specifically, how differences in cellular mechanics at the interface between oncogenic and wild-type populations influence tumor initiation and progression is poorly understood. Addressing this gap could not only enhance our understanding of cancer biology but also open new avenues for tissue-specific therapeutic interventions by identifying what makes certain tissues more permissive to tumor formation.

To this end, we investigated and compared the bio-physical factors influencing early tumorigenesis in human mammary and bronchial epithelial monolayers by sporadically transfecting a subset of healthy cells into oncogenic cells expressing a constitutively active HRasV12 protein. We then tracked the fate of these transfected cells and their interactions with the wild-type counterparts in real time via live imaging. Strikingly, HRasV12-mutant cells exhibited contrasting behaviors in the two tissues. In mammary epithelium, isolated transformed cells were frequently extruded. However, when present in groups, extrusion of oncogenic cells was not seen, but instead, they remained spatially confined in a jammed state. In contrast, in bronchial epithelium, single as well as groups of HRasV12 expressing cells persisted and exhibited unrestricted expansion over time, regardless of cluster size. To explain these findings, we employed a bi-disperse vertex model of epithelia [18-21] and demonstrated that differences in interfacial tension between wild-type and transformed populations play a crucial role in determining whether mutant cells are eliminated, restrained, or are allowed to persist and expand. Taken together, these findings provide novel insights into the differential outcomes of competitive cellular interactions in the precancer stages and uncover a mechanical basis for tissue-specific oncogenesis.

Results

Singlets and groups of HRasV12 oncogenic mutants show tissue-specific fates in epithelia

To explore the differences in the premalignant stages across epithelial tissues, we sporadically transfected monolayers of the two non-transformed human epithelial cell lines, MCF10A (mammary) and BEAS2B (bronchial), with a plasmid construct that constitutively expresses HRasV12 along with the GFP reporter [Figure 1a]. We then monitored the cellular dynamics over 90 hours. Consistent with previous studies reporting extrusion of single HRasV12 mutants in MDCK monolayers [11-15, 22], we observed that isolated HRasV12 transfected cells (singlets) surrounded by wild-type neighbours were extruded from the mammary epithelium [Figure. 1b (upper panel), Supplementary video 1]. This process correlated with a reduction in shape indices of oncogenic mutants as the density of wild-type cells around them increased [Figure. 1c (upper panel)]. Interestingly, this well-known mechanism of epithelial defense against cancer was impaired in BEAS2B cells, where HRasV12 mutants evaded extrusion [Figure. 1b (lower panel), d], and wild-type cell compaction did not restrict their growth [Figure. 1c (lower panel), Supplementary video 1]. Instead, the mutants in the bronchial epithelium continued to proliferate and exhibited increased elongation with long protrusions [Figure. 1c (lower panel), e, Supplementary video 1]. Given that cancers frequently arise from fields of oncogenic mutants rather than isolated mutated cells [23-25], we also investigated whether these differential dynamics of mutant singlets would also extend to groups of oncogenic mutants, potentially enabling more aggressive HRasV12 growth in bronchial tissue compared to mammary tissue. By adjusting cell seeding conditions, we generated larger groups of HRasV12 transfected cells in both MCF10A and BEAS2B monolayers and then tracked the dynamics of these groups over time. In the mammary epithelium, wild-type cells were unable to extrude groups of HRasV12-transfected cells but instead confined them into compact, circular clusters with smooth, rounded interfaces [Figure 1f, (upper panel), g (upper panel), Supplementary Video 2]. As a result, oncogenic clusters gradually demixed from the surrounding wild-type MCF10A population and remained spatially constrained [Figure 1h]. In contrast, HRasV12 clusters in the bronchial epithelium gradually spread outward [Figure 1f, (lower panel), g (lower panel), Supplementary video 2], and formed protrusive lamellipodia and developed an irregular interface with surrounding wild-type cells [Figure. 1g (lower panel)]. Together, these striking differences in the dynamics of HRasV12 oncogenic cells demonstrate that activation of the same oncogene can lead to distinct outcomes across tissues in terms of growth, spreading, and extrusion. This prompted us to further investigate the unique bio-physical mechanisms driving these variations.

Tissue-specific outcomes of HRasV12 oncogenic mutants in epithelial monolayers

(a) Schematic representation of the experimental setup and outcome (b) Representative images showing extrusion of single oncogenic cells (singlets) in mammary epithelium (upper panel), but protrusive behavior of the same oncogenic singlets in bronchial epithelium (bottom panel) (c) Quantification of shape indices of HRasV12 singlets showing a reduction in shape indices, with increasing wild-type cell density in mammary epithelium (upper panel), but the opposite trend in bronchial epithelium (bottom panel), where HRasV12 cells continues to spread. Shaded regions indicate the results simple linear regression analysis with 95% confidence intervals (d) Extrusion rates of HRasV12 singlets are significantly higher in mammary epithelium, while (e) Percentage of HRasV12 cells forming protrusions is markedly higher in bronchial epithelium. Representative data are plotted from one of three independent experiments with the median shown as a bold dashed line and the first and third quartiles are shown as thin dashed lines. Statistical significance was calculated using Unpaired t-test with Welch’s correction. (f) Behaviour of HRasV12 oncogenic clusters in the two tissues-In mammary epithelium (upper panel), clusters become spatially confined with a smooth, circular interface with the wild-type population, while in bronchial epithelium (bottom panel), HRasV12 clusters expand, forming long protrusions over time (g) Immunofluorescence images of HRasV12 clusters in mammary and bronchial epithelia, highlighting differences in spreading patterns in the two epithelia (h) Quantification of cluster segregation: mammary epithelium exhibits a higher segregation index and reduced eccentricity compared to bronchial epithelium, indicating greater spatial confinement of oncogenic clusters. The segregation index (SI), defined as the average ratio of homotypic and all cell neighbors, quantifies the degree of demixing and is averaged over all cells inside an oncogenic cluster. Data are mean±sem and plotted from different clusters from one representative experiment. (Scale bars= 50 μm)

Biophysical analysis reveals demixing via jamming of oncogenic clusters in MCF10A but unjamming and protrusive growth in BEAS2B

To understand the mechanisms driving the differential fates of oncogenic fields in the two epithelial tissues, we mapped the bio-physical signatures of wild-type-mutant interactions by subdividing the images into four regions of interests for both monolayers-ROI1: wild type cells distant from oncogenic cells, ROI2: interfacial wild type cells in direct contact with oncogenic cells, ROI3: oncogenic cells at the cluster boundaries in direct contact with wild type cells, and ROI4: oncogenic cells in contact with other oncogenic cells [Figure. 2a]. We then compared the F-actin levels and shapes of both mutant and wild-type populations across these ROIs. Additionally, we employed Bayesian force inference, as previously established [26], to infer the differential tissue stresses associated with cluster formation from the observed cell geometries. Interestingly, the differences across these parameters were apparent not only between wild-type and oncogenic counterparts but also based on the spatial localization of the cells, with cells at the oncogenic-wild type interface exhibiting distinct behaviors, regardless of their origin.

Distinct mechanical signatures of HRasV12 clusters in mammary and bronchial epithelia

(a) Representative images showing regions of interest (ROIs) selected for analysis in mammary (upper panel) and bronchial epithelia (lower panel) (b) Immunofluorescence images of HRasV12 clusters stained for F-actin in mammary (upper panel) and bronchial epithelia (lower panel) (c) Oncogenic cluster–wild-type interface showing distinct actin belt in mammary marked by white arrows (upper panel) and absence of it in bronchial epithelia (lower panel) (d) ROI-based F-actin stained regions of wild-type and oncogenic cells in mammary (upper panel) and bronchial epithelia (lower panel) revealing distinct shape differences in different regions-in mammary epithelia, wild type cells at the interface show elongated shapes (ROI2), and oncogenic cells inside the cluster show jammed shapes (ROI4), in comparison to wild type cells at a random location (ROI1), while in bronchial epithelia, oncogenic cells (ROI 3 and 4) show more elongated shapes compared to wild-type beas2b cells (e), (f) Quantifications of F-actin intensities and shape indices, in mammary epithelia (e), and bronchial epithelia (f), across the four ROIs confirming jammed oncogenic clusters, with unjamming at the interface, in mammary epithelia and unjammed clusters in bronchial epithelia (g) Schematic illustrating the Bayesian Force Inference pipeline used to estimate relative intracellular pressure and cell–cell edge tension (h) Heatmaps depicting relative cell pressures and cell–cell edge tensions in mammary epithelium, revealing localized mechanical heterogeneity (i) Heatmaps of relative cell pressures and edge tensions in bronchial epithelium, showing a distinct mechanical landscape compared to mammary epithelium (j) Quantification of relative intracellular pressures and cell–cell edge tensions across the four ROIs in mammary (upper panel) and bronchial epithelia (lower panel) highlighting tissue-specific differences. All data represented as mean±sem in box and whisker plots and are plotted from one out of three independent experiments with whiskers extending from the minimum to the maximum values, showing the full range of the dataset including outliers. Line represents the median. Statistical significance was calculated using Unpaired t-test with Welch’s correction. Scale bars= 20 μm

Earlier version: In mammary epithelium, oncogenic clusters were surrounded by a prominent actin belt [Figure. 2b (upper panel), c (upper panel)], coupled with drastically low F-actin levels inside the clusters [Figure. 2b (upper panel), e (left panel)], suggesting that the de-mixing of oncogenic cells in mammary tissue is likely driven by changes in cytoskeletal mechanics. Further, oncogenic clusters in mammary epithelium (ROI4) also exhibited low cell shape indices [Figure. 2d (upper panel), e (right panel)], high internal pressure and low line tension [Figure. 2h, j]-characteristics indicative of jamming [27]. In epithelial tissues, the jamming–unjamming transition is closely linked to changes in cell shape [27] and is often quantified using the cell shape index (the ratio of cell perimeter to the square root of area), where lower shape indices correspond to jammed, solid-like states, and higher shape indices reflect unjammed, fluid-like tissue behavior. Notably, both oncogenic cells and wildtype MCF10A cells at the cluster interface (ROI 2 and ROI3) showed F-actin enrichment, while the interfacial wild shape cells show elongated shapes [Figure. 2d (upper panel), e], indicating that interfacial cellular mechanics play a key role in the de-mixing process as the wild-type cells encircle the oncogenic mutants. This interfacial region enriched with F-actin also contained phosphomyosin [Supplementary Figure. 1a], indicating that a contractile actomyosin belt drives the compaction of oncogenic clusters in mammary epithelia. In contrast, bronchial epithelium exhibited no actin belt formation at the wild type-oncogenic cell interface, and no traits of jamming were observed in the oncogenic clusters [Figure. 2b (lower panel), c (lower panel)]. Instead, oncogenic cells continued to proliferate, displaying elongation measured via higher shape indices and high F-actin expression [Figure. 2d (lower panel), f]-indicative of unjamming behaviour [27,28]. The relative internal cell pressure and cell-cell line tension inferred with Bayesian force inference showed slight or no modulation in the oncogenic clusters, given the general elongated shapes of wild-type cells in BEAS2B cells [Figure. 2i, j]. Together, these findings highlight distinct interfacial mechanics between oncogenic cells and their wild-type neighbors in both MCF10A and BEAS2B monolayers, resulting in differing outcomes in mutant colony growth.

Revised version: In mammary epithelium, oncogenic clusters were surrounded by a contractile actomyosin belt [Figure. 2b (upper panel), c (upper panel)], shown by a prominent deposition of F-actin and phosphomyosin at heterotypic (oncogenic cluster-wild type population) interfaces [Supplementary Figure 4] coupled with drastically low F-actin levels inside the clusters [Figure. 2b (upper panel), e (left panel)], suggesting that the de-mixing of oncogenic cells in mammary tissue is likely driven by changes in cytoskeletal mechanics. Further, oncogenic clusters in mammary epithelium (ROI4) also exhibited low cell shape indices [Figure. 2d (upper panel), e (right panel)], high internal pressure and low line tension [Figure. 2h, j]-characteristics indicative of jamming [27]. In epithelial tissues, the jamming–unjamming transition is closely linked to changes in cell shape [27] and is often quantified using the cell shape index (the ratio of cell perimeter to the square root of area), where lower shape indices correspond to jammed, solid-like states, and higher shape indices reflect unjammed, fluid-like tissue behavior. Notably, both oncogenic cells and wildtype MCF10A cells at the cluster interface (ROI 2 and ROI3) showed F-actin enrichment, while the interfacial wild shape cells show elongated shapes [Figure. 2d (upper panel), e], indicating that interfacial cellular mechanics play a key role in the de-mixing process as the wild-type cells encircle the oncogenic mutants. This interfacial region enriched with F-actin also contained phosphomyosin [Supplementary Figure. 1a], indicating that a contractile actomyosin belt drives the compaction of oncogenic clusters in mammary epithelia. In contrast, bronchial epithelium exhibited no actin belt formation at the wild type-oncogenic cell interface, and no traits of jamming were observed in the oncogenic clusters [Figure. 2b (lower panel), c (lower panel)]. Instead, oncogenic cells continued to proliferate, displaying elongation measured via higher shape indices and high F-actin expression [Figure. 2d (lower panel), f]-indicative of unjamming behaviour [27,28]. The relative internal cell pressure and cell-cell line tension inferred with Bayesian force inference showed slight or no modulation in the oncogenic clusters, given the general elongated shapes of wild-type cells in BEAS2B cells [Figure. 2i, j]. Together, these findings highlight distinct interfacial mechanics between oncogenic cells and their wild-type neighbors in both MCF10A and BEAS2B monolayers, resulting in differing outcomes in mutant colony growth.

PIV analysis reveals kinetic arrest of oncogenic clusters in MCF10A coupled with tangential motion of interfacial cells

Earlier version: To further probe the cellular dynamics that lead to the contrasting fates of the oncogenic mutant colonies in the two tissues, we analyzed cellular velocities during the de-mixing of oncogenic cells, using Particle Image Velocimetry (PIV) on consecutive image pairs from mammary and bronchial monolayers. Interestingly, velocity maps in samples imaged for mutant cluster – wild type dynamics in MCF10A [Supplementary video 3] revealed an unusual movement pattern at confluency: wild type MCF10A cells exhibited tangential motion around oncogenic clusters [Figure. 3c], which coupled with a drastic reduction in the velocity of oncogenic cells [Figure. 3a (lower panel), d], likely caused their compaction and jamming into clusters. Notably, this tangential movement was unique to wild-type MCF10A cells surrounding oncogenic clusters and was absent in other regions across the monolayer [Figure. 3c] as well as under control conditions-i.e., in the absence of oncogenic mutants [Supplementary video 4], in both of which a typical velocity drop was observed in MCF10A cells due to density-induced jamming [27-30] [Figure. 3d (upper panel), Supplementary 2]. These results aligned with the observed differences in cell shape indices, where oncogenic clusters underwent jamming, while wild-type cells at the interface exhibited unjamming behavior [29,30]. In contrast, in the bronchial epithelium, mutant clusters remained unjammed, even as surrounding wild-type cells became more crowded [Figure. 3b, d (lower panel), Supplementary video 5]. Additionally, no tangential cellular movement was detected around the oncogenic clusters in the bronchial epithelium. These velocity maps showed negligible deviations from wild type (control) BEAS2B velocity maps [Supplementary video 6]. Together, these findings suggest that mutant clusters carrying the same oncogene can exhibit drastically different behaviors depending on the epithelial tissue type. Specifically, differences in the dynamics of the interfacial cells drive the demixing and jamming of HRasV12 oncogenic cells in mammary epithelium while promoting protrusive, unjammed growth in bronchial epithelium.

Revised version: To further probe the cellular dynamics that lead to the contrasting fates of the oncogenic mutant colonies in the two tissues, we analyzed cellular velocities during the de-mixing of oncogenic cells, using Particle Image Velocimetry (PIV) on consecutive image pairs from mammary and bronchial monolayers. Interestingly, velocity maps in samples imaged for mutant cluster – wild type dynamics in MCF10A [Supplementary video 3] revealed an unusual movement pattern at confluency: wild type MCF10A cells exhibited tangential motion around oncogenic clusters [Figure. 3c], which coupled with a drastic reduction in the velocity of oncogenic cells [Figure. 3a (lower panel), d], likely caused their compaction and jamming into clusters. Notably, this tangential movement was unique to wild-type MCF10A cells surrounding oncogenic clusters and was absent in other regions across the monolayer [Figure. 3c, Supplementary Figure 5] as well as under control conditions-i.e., in the absence of oncogenic mutants [Supplementary video 4], in both of which a typical velocity drop was observed in MCF10A cells due to density-induced jamming [27-30] [Figure. 3d (upper panel), Supplementary Figure 2]. These results aligned with the observed differences in cell shape indices, where oncogenic clusters underwent jamming, while wild-type cells at the interface exhibited unjamming behavior [29,30]. In contrast, in the bronchial epithelium, mutant clusters remained unjammed, even as surrounding wild-type cells became more crowded [Figure. 3b, d (lower panel), Supplementary video 5]. Additionally, no tangential cellular movement was detected around the oncogenic clusters in the bronchial epithelium. These velocity maps showed negligible deviations from wild type (control) BEAS2B velocity maps [Supplementary video 6]. Together, these findings suggest that mutant clusters carrying the same oncogene can exhibit drastically different behaviors depending on the epithelial tissue type. Specifically, differences in the dynamics of the interfacial cells drive the demixing and jamming of HRasV12 oncogenic cells in mammary epithelium while promoting protrusive, unjammed growth in bronchial epithelium.

Particle Image Velocimetry (PIV) reveals distinct cellular movement patterns in HRasV12 clusters in the two monolayers

(a), (b) Representative snapshots of HRasV12 clusters, and the corresponding PIV maps, in mammary (a) and bronchial epithelium (b), as the density of wild type cells increase showing wild type cells along the cluster-wild-type interface exhibit a higher velocity attributed to tangential motion in mammary epithelium (a), but undergo a more kinetically arrested state in bronchial epithelium (b) (c) Oncogenic cluster in MCF10A monolayer with concentric circles drawn from the centre of the cluster, radially outward till its boundary (upper panel) with the red (highlighted) circle representing the one along which tangential velocity was the highest and plot showing tangential motion (expressed as a function of the circles), revealing highest tangential motion (data point highlighted with red) along the interface of oncogenic cluster with the wild-type cells in mammary epithelia (lower panel). No significant tangential motion in random wild-type region (c) Representative snapshots of HRasV12 clusters in bronchial epithelium at the same two time points and the corresponding PIV velocity maps (d) Quantification of cellular velocities in local regions as cell densities reach confluency showing oncogenic cells clustered in mammary epithelium (top panel) jam and show a significant reduction in velocities. In bronchial epithelium (bottom panel) while wild type cells show reduced movement as cell density increased, transfected cells continued to show higher motion, signifying that jamming of wild type cells has no effect on mutants. All data are represented as mean±sem plotted from different clusters from one representative experiment. Scale bars= 50 μm.

Earlier version: Bi-disperse vertex model reveals differential interfacial tension between mutant and wild-type population governs tissue-specific response towards oncogenic mutants

Revised version: Bi-disperse vertex model reveals interfacial tension between mutant and wild-type population governs tissue-specific response towards oncogenic mutants

To explain the differential dynamics and degrees of demixing of oncogenic clusters observed in the two tissues, we drew inspiration from the differential interfacial tension hypothesis (DITH) [31]. This hypothesis suggests that cell sorting in tissues could occur due to the variations in interfacial tension, as seen in embryonic cell rearrangement, where cells regulate contractility, as well as adhesions to demix from their neighbors [31-33]. To test whether such a mechanism could be applied to the de-mixing of oncogenic clusters from wild-type cells-leading to tissue-specific outcomes, we employed the cell-based bi-disperse Vertex Model (VM). Our model allowed us to simulate an epithelial monolayer and capture the dynamic interactions between wild-type and mutant cell populations. We analyzed the mechanics underlying the segregation dynamics of mutant clusters by examining the interfacial line tension ΛijM/WT between the oncogenic and wild-type cells [Figure. 4a]. In our model, Λij = 0, for homotypic interfaces, i.e. wild-type-wild type or mutant-mutant interactions, and Λij = Λ ≠ 0, for heterotypic interfaces between oncogenic and wild type cells. When Λ > 0, our model ensures that the interface between oncogenic and healthy cells contracts to minimize energy, leading to a reduction in total interfacial length, and when Λ < 0, our model promotes a rough interface by favouring increased interface between healthy and oncogenic cells [Supplementary Video 7]. The final interface shape is determined by the balance between interfacial tension and the intrinsic perimeter and area elasticity of individual cells [Figure. 4b].

Bi-disperse vertex model reveals interfacial tension-driven segregation of oncogenic clusters

(a) Schematic of the bi-disperse vertex model, where interfacial line tension (red) is applied at the boundary between mutant (green) and wild-type (grey) (b) The value of Λ determines whether the mutant cluster remains compact or spreads into the wild-type population, where the red line gives the outline of the morphology of the mutant cluster interface (c) Difference in shape index between interfacial and bulk cells as a function of Λ (interfacial tension), shown for different stress thresholds. Experimental average (red) overlaid for comparison (d) Eccentricity of mutant clusters (given by the red lines in (b)) over time for different values of Λ (with Π0 = 0.3), indicating changes in cluster morphology (e) Segregation index at t = 300 as a function of Λ forΠ0 = 0.3, quantifying the extent of cluster segregation (f) Time evolution of the segregation index for different values of Λ (with Π0 = 0.3), demonstrating interfacial tension-driven segregation (g) Hydrostatic pressure maps of mutant and wild-type cells at t = 300, with mutant cells labeled by green directors and wild-type cells by black directors (h) Velocity field maps of epithelial monolayers for different values of Λ, revealing flow patterns around oncogenic clusters (i) Tangential velocity loop integral for circles of varying radius centered at the mutant cluster (Region 1) and centered away from the cluster (Region 2), showing tissue-specific velocity variations (j) Extrusion probability of a single mutant cell as a function of Λ, demonstrating how interfacial tension influences mutant cell fate.

To explain the observed differences in interfacial cellular dynamics, we introduced an active contribution to interfacial tension, which was coupled to the shape of cells at the interface, called shape-tension coupling[34]. This active interfacial tension is extensile, meaning it elongates cells along the interface depending on their orientation relative to the interface. It originates from the anisotropic distributions of stress-fibres in when cells are under external stress. Hence, we assume that it becomes significant only when the hydrostatic pressure on the interfacial cells exceeds a threshold value Π0 [Figure. 4b, g].

Remarkably, our model robustly captured the experimentally observed behaviors. As a first check, we compared the experimentally observed differences in the shape indices between interfacial and non-interfacial wild-type and oncogenic cells in both mammary and bronchial epithelium with our model’s predictions. For Γ = 0.2, Π0 = {0, 0.1, 0.3}, Λ = {0 − 1.6}, our model accurately captured the experimental observations [Figure. 4c]. For stress threshold values lower than Π0 = 0.3, we find that the differences in shape indices are non-zero even for Λ = 0, which is an unphysical situation. This is unphysical because we expect there to be no difference in the shape indices because in principle, there is only one cell type for Λ = 0. Thus, we have taken Π0 = 0.3 for further analyses. To make the comparison between the model and the experiment more robust, we compared other metrics as well.

Earlier version: To quantify the extent of mixing between wild-type and mutant cells, we used the segregation index. Our simulations showed that when Λ < 0 the segregation index rapidly decreases over time, indicating outward protrusion of the mutant cluster. In contrast, for Λ > 0 the segregation index remained close to 1, reflecting a compact, jammed cluster [Figure. 4f]. Additionally, we observed that when Λ > 0, cells within the cluster experience higher internal pressure compared to when Λ < 0 [Figure 4g], consistent with experimental findings [Figure. 2h,i]. Furthermore, flow field analysis from our simulations [Figure 4h,i] revealed that for Λ > 0, cell movement aligns tangentially along the cluster boundary—consistent with the MCF10A behaviour—whereas for Λ < 0 cell motion was uniform throughout the monolayer, resembling the BEAS2B phenotype [Figure. 3a, b]. These comparions further establish that different values of Λ correspond to different cell types. The same oncogene thus might affect different cell lines differently. We also estimated the probability of single mutant cell extrusion as a function of Λ [Figure. 4j], consistent with the findings from our experiments which show that isolated HRasV12 transfected cells are frequently extruded in mammary tissue [Figure. 1b]. Further, to check how actin distribution in the tissue can change in response to the interfacial tension in the model, we performed simulations of a vertex model with interfacial tension and mechanochemical feedback [35]. The simulation snapshots matched quite well with those from experiment [Appendix figure 5]. Together, our findings demonstrate how interfacial tension-driven mechanics could play a crucial role in determining the competitive dynamics between wild-type and mutant populations, thus providing compelling evidence that cellular and tissue mechanics are key determinants of cell competition.

Revised version: To quantify the extent of mixing between wild-type and mutant cells, we used the segregation index. Our simulations showed that when Λ < 0 the segregation index rapidly decreases over time, indicating outward protrusion of the mutant cluster. In contrast, for Λ > 0 the segregation index remained close to 1, reflecting a compact, jammed cluster [Figure. 4f]. Additionally, we observed that when Λ > 0, cells within the cluster experience higher internal pressure compared to when Λ < 0 [Figure 4g], consistent with experimental findings [Figure. 2h,i]. Line tensions along edges in the vertex model have also been calculated [Supplementary Figure 6]. Furthermore, flow field analysis from our simulations [Figure 4h,i] revealed that for Λ > 0, cell movement aligns tangentially along the cluster boundary—consistent with the MCF10A behaviour—whereas for Λ < 0 cell motion was uniform throughout the monolayer, resembling the BEAS2B phenotype [Figure. 3a, b]. These comparions further establish that different values of Λ correspond to different cell types. The same oncogene thus might affect different cell lines differently. We also estimated the probability of single mutant cell extrusion as a function of Λ [Figure. 4j], consistent with the findings from our experiments which show that isolated HRasV12 transfected cells are frequently extruded in mammary tissue [Figure. 1b]. Further, to check how actin distribution in the tissue can change in response to the interfacial tension in the model, we performed simulations of a vertex model with interfacial tension and mechanochemical feedback [35]. The mechanochemical feedback incorporates a load-dependent myosin-binding mechanism, allowing cells in the vertex model to adjust their fractions of stress fibers and cortical actin individually. The simulation snapshots matched quite well with those from experiment [Appendix figure 5]. Together, our findings demonstrate how interfacial tension-driven mechanics could play a crucial role in determining the competitive dynamics between wild-type and mutant populations, thus providing compelling evidence that cellular and tissue mechanics are key determinants of cell competition.

Actomyosin belt disruption prevents mutant cluster compaction in mammary epithelia

(a) Representative actin staining image upon blebbistatin treatment, done after oncogenic induction in mammary epithelial monolayer, showing an absence of an actin belt (b) Quantification of F-actin intensities (left) and shape indices (right), showing difference between oncogenic clusters and wild-type cells in these parameters, is getting small – although not completely vanished because of low concentration of blebbistatin used for the experiment (c) Representative actin and E-cadherin staining images of the wild type monolayers showing shorter stress fibers, higher E-cadherin intensities in MCF10A wild-type monolayer, and the opposite in BEAS2B (d) Plots comparing F-actin stress fiber length (left panel) and shape indices (right panel) in wild-type MCF10A and BEAS2B monolayers. Stress fiber lengths were calculated manually in ImageJ (e) Plots comparing F-actin and E-cadherin levels in the two wild-type tissues. Representative data are plotted from one of three independent experiments with the median shown as a bold line. Statistical significance was calculated using Unpaired t-test with Welch’s correction. Scale bars= 50 μm.

Disrupting actomyosin belt using blebbistatin prevents demixing and compaction of mutant clusters in mammary epithelia

Added in the revised version: Since our model showed that differences in interfacial tension regulate mutant cluster dynamics, with a higher interfacial tension driving cluster compaction [Supplementary video 7], we reasoned that this tension is mediated by the supracellular actomyosin belt [36] that drives the compaction and segregation of oncogenic clusters in mammary epithelium. Therefore, we investigated whether disrupting the cytoskeletal organization and impairing actomyosin belt formation would affect the demixing and compaction of oncogenic clusters. To test this, we treated MCF10A monolayers with 5 µM blebbistatin, a myosin II inhibitor, to disrupt belt formation.

Added in the revised version: As expected, blebbistatin treatment prevented mutant cluster compaction [Figure. 5a], leading to a loss of their rounded morphology and restoring actin levels and shape indices of mutant cells to wild-type levels [Figure. 5b]. This shift resulted in mutant clusters resembling those observed in bronchial epithelia, suggesting that cytoskeletal organization plays a key role in tissue-specific differences in mutant cluster dynamics. This was also consistent with the fact that interfacial tension between cells arises from the interplay between the cortical actomyosin network and cell-cell adhesion [37].

Earlier version: Since our model showed that differences in interfacial tension regulate mutant cluster dynamics, with a higher interfacial tension driving cluster compaction [Supplementary video 7], we reasoned that this tension is mediated by the supracellular actomyosin belt [36] that drives the compaction and segregation of oncogenic clusters in mammary epithelium. Therefore, we investigated whether disrupting the cytoskeletal organization and impairing actomyosin belt formation would affect the demixing and compaction of oncogenic clusters. To test this, we treated MCF10A monolayers with 5 µM blebbistatin, a myosin II inhibitor, to disrupt belt formation. As expected, blebbistatin treatment prevented mutant cluster compaction [Figure. 5a], leading to a loss of their rounded morphology and restoring actin levels and shape indices of mutant cells to wild-type levels [Figure. 5b]. This shift resulted in mutant clusters resembling those observed in bronchial epithelia, suggesting that cytoskeletal organization plays a key role in tissue-specific differences in mutant cluster dynamics. This was also consistent with the fact that interfacial tension between cells arises from the interplay between the cortical actomyosin network and cell-cell adhesion [37]. Additionally, we looked into the inherent differences in cytoskeletal mechanics of the host tissues and observed distinct differences in actin organization between the two tissues. Wild-type bronchial and mammary epithelial cells exhibited comparable F-actin intensities [Figure. 5c (upper panel), e], although the bronchial epithelial cells had longer stress fibers [Figure. 5c (middle panel), d (left panel)]. These cytoskeletal differences were also coupled with significant differences in cell shapes between confluent MCF10A and BEAS2B monolayers with wild-type BEAS2B cells exhibited significantly higher shape indices, indicating elongated, loosely packed morphologies and a lower degree of jamming [Figure 5d (right panel). In contrast, mammary epithelial cells had lower shape indices, reflecting a more compact and tightly packed arrangement [Figure. 5d (right panel)]. These differences were also in agreement with the PIV maps of the two wild-type tissues [Supplementary figure 2, Supplementary videos 4 and 6] where the mammary monolayers underwent a sharper drop in velocities and therefore kinetic arrest, compared to the bronchial epithelia. Bronchial epithelial cells also exhibited weaker cell-cell adhesions, as indicated by reduced E-cadherin staining intensity at cell junctions compared to mammary epithelial cells [Figure. 5c (bottom panel), e]. These weaker adhesions likely contributed to increased cellular motility and a diminished ability to constrain mutant clusters, in contrast to the more stable, compacted clusters observed in mammary epithelium. Together, these findings demonstrate that variations in mechanical properties-including cell shape, actin organization, and adhesion strength-govern the differential responses of mammary and bronchial epithelia to oncogenic mutants and may in turn, influence tumor growth and progression.

Revised version: Additionally, we looked into the inherent differences in cytoskeletal mechanics of the host tissues and observed distinct differences in actin organization between the two tissues. Wild-type bronchial and mammary epithelial cells exhibited comparable F-actin intensities [Figure. 5c (upper panel), e], although the bronchial epithelial cells had longer stress fibers [Figure. 5c (middle panel), d (left panel)]. These cytoskeletal differences were also coupled with significant differences in cell shapes between confluent MCF10A and BEAS2B monolayers with wild-type BEAS2B cells exhibited significantly higher shape indices, indicating elongated, loosely packed morphologies and a lower degree of jamming [Figure 5d (right panel). In contrast, mammary epithelial cells had lower shape indices, reflecting a more compact and tightly packed arrangement [Figure. 5d (right panel)]. These differences were also in agreement with the PIV maps of the two wild-type tissues [Supplementary figure 2, Supplementary videos 4 and 6] where the mammary monolayers underwent a sharper drop in velocities and therefore kinetic arrest, compared to the bronchial epithelia. Bronchial epithelial cells also exhibited weaker cell-cell adhesions, as indicated by reduced E-cadherin staining intensity at cell junctions compared to mammary epithelial cells [Figure. 5c (bottom panel), e]. These weaker adhesions likely contributed to increased cellular motility and a diminished ability to constrain mutant clusters, in contrast to the more stable, compacted clusters observed in mammary epithelium. Together, these findings demonstrate that variations in mechanical properties-including cell shape, actin organization, and adhesion strength-govern the differential responses of mammary and bronchial epithelia to oncogenic mutants and may in turn, influence tumor growth and progression.

Discussion

Earlier version: A central challenge in cancer biology is understanding why some oncogenic mutations are more potent in certain cancers. Large-scale cancer genomics has revealed that many driver mutations are shared across epithelial cancers, yet tumor incidence, progression rates, and invasive potential vary widely between organs [38, 39]. These observations have traditionally been explained through tissue-specific differences in signaling pathways, differentiation states, and mutational landscapes. However, such explanations are incomplete, particularly in the context of early tumorigenesis, where oncogenic mutations are often detected in histologically normal tissues without leading to overt cancer [38]. This raises a fundamental question: what prevents or permits oncogenic cells from expanding at the earliest stages of cancer initiation? The present study addresses this gap by proposing interfacial mechanics between wild-type and nascently transformed cells, as a central organizing principle in early cancer initiation. Rather than focusing oncogenic cells in isolation, our work frames tumor initiation as an emergent outcome of mutant–wild-type interactions within an epithelial collective. From this perspective, the critical determinant of oncogenic fate is not only the presence of a mutation, but how the host tissue mechanically responds to it. This conceptual shift aligns with growing recognition that epithelial tissues behave as active materials, whose collective mechanical state can either suppress or amplify local perturbations. These insights shed light on the differential susceptibility of epithelial tissues to oncogenic transformation and uncover a mechanical basis for tissue-specific oncogenesis. By demonstrating how interfacial tension at the mutant-wild-type boundary governs whether oncogenic cells are restrained or overgrow, our work shifts the focus from purely molecular pathways to the biophysical constraints that regulate tumor initiation. While previous studies have emphasized the contextual nature of tumorigenesis [2-4] and highlighted the critical role of interactions between healthy and oncogenic cells during epithelial cancer initiation [14, 22, 40], our work integrates and extends these perspectives by demonstrating how interfacial tension at the mutant-wild-type boundary determines whether transformed cells remain constrained or undergo unchecked expansion.

Revised version: A central challenge in cancer biology is understanding why some oncogenic mutations are more potent in certain cancers. Large-scale cancer genomics has revealed that many driver mutations are shared across epithelial cancers, yet tumor incidence, progression rates, and invasive potential vary widely between organs [38, 39]. These observations have traditionally been explained through tissue-specific differences in signaling pathways, differentiation states, and mutational landscapes. However, such explanations are incomplete, particularly in the context of early tumorigenesis, where oncogenic mutations are often detected in histologically normal tissues without leading to overt cancer [38]. This raises a fundamental question: what prevents or permits oncogenic cells from expanding at the earliest stages of cancer initiation?

Added in the revised version: The present study addresses this gap by proposing interfacial mechanics between wild-type and nascently transformed cells, as a central organizing principle in early cancer initiation. Rather than focusing oncogenic cells in isolation, our work frames tumor initiation as an emergent outcome of mutant–wild-type interactions within an epithelial collective. From this perspective, the critical determinant of oncogenic fate is not only the presence of a mutation, but how the host tissue mechanically responds to it. This conceptual shift aligns with growing recognition that epithelial tissues behave as active materials, whose collective mechanical state can either suppress or amplify local perturbations. These insights shed light on the differential susceptibility of epithelial tissues to oncogenic transformation and uncover a mechanical basis for tissue-specific oncogenesis. By demonstrating how interfacial tension at the mutant-wild-type boundary governs whether oncogenic cells are restrained or overgrow, our work shifts the focus from purely molecular pathways to the biophysical constraints that regulate tumor initiation.

Added in the revised version: While previous studies have emphasized the contextual nature of tumorigenesis [2-4] and highlighted the critical role of interactions between healthy and oncogenic cells during epithelial cancer initiation [14, 22, 40], our work integrates and extends these perspectives by demonstrating how interfacial tension at the mutant-wild-type boundary determines whether transformed cells remain constrained or undergo unchecked expansion.

Earlier version: To probe the mechanistic basis of these differences, we applied the Differential Interfacial Tension Hypothesis (DITH) [31,32] and demonstrated that interfacial tension is a key determinant of mutant cluster morphology and dynamics. When interfacial tension between wild-type and mutant populations was positive, mutant clusters shrunk into circular aggregates and stayed constrained. In contrast, a negative interfacial tension resulted in irregular, protrusive clusters that failed to compact and kept growing. Importantly, while the Differential Interfacial Tension Hypothesis poses that segregation results from global differences in the contractility and adhesive properties of two tissues displaying different intrinsic tensions, the results of the present work support a different scenario, where what counts is the difference in the heterotypic interfacial tension, i.e, the tension along the boundary of the wild-type and mutant populations. In line with this, we also showed that disrupting the interfacial actomyosin belt around the mutant clusters in the mammary epithelia, disrupted the restraint on them, preventing their demixing and compaction and shifting cluster behavior to a more unjammed state, similar to the ones in the bronchial system. These findings underscore the role of interfacial cellular mechanics in shaping oncogenic cell fate, with mammary epithelium imposing physical constraints that limit mutant expansion, while bronchial epithelium lacks such mechanical barriers, permitting unchecked growth.

Revised version: Importantly, while classical frameworks based on interfacial tension have often been interpreted through the lens of the Differential Interfacial Tension Hypothesis (DITH) [31,32] which attributes segregation to global differences in the contractility and adhesive properties of two tissues displaying different intrinsic tensions, the results of the present work support a different scenario, where what counts is the difference in the heterotypic interfacial tension, i.e, the tension along the boundary of the wild-type and mutant populations. We demonstrate that this heterotypic interfacial tension is a key determinant of mutant cluster morphology and dynamics. This is also consistent with recent theoretical work showing that sharp interfaces can emerge from interfacial mechanics alone [41]. When interfacial tension between wild-type and mutant populations was positive, mutant clusters shrunk into circular aggregates and stayed constrained. In contrast, a negative interfacial tension resulted in irregular, protrusive clusters that failed to compact and kept growing. In line with this, we also showed that disrupting the interfacial actomyosin belt around the mutant clusters in the mammary epithelia, disrupted the restraint on them, preventing their demixing and compaction and shifting cluster behavior to a more unjammed state, similar to the ones in the bronchial system. These findings underscore the role of interfacial cellular mechanics in shaping oncogenic cell fate, with mammary epithelium imposing physical constraints that limit mutant expansion, while bronchial epithelium lacks such mechanical barriers, permitting unchecked growth.

By integrating mechanical constraints into models of cancer initiation, our study provides a framework for understanding how different epithelial tissues resist or permit tumorigenesis. However, while interfacial tension emerges as a key regulator of oncogenic segregation, the molecular pathways driving these mechanical differences remain to be elucidated. Future studies should investigate how targeting specific signaling networks or cytoskeletal components influences interfacial tension and whether modulating these factors could provide therapeutic benefits. Additionally, extending this framework to other oncogenes and tissue types could further elucidate the mechanical principles governing tumor initiation and progression.

In conclusion, this work supports a growing view of cancer initiation as a mechanically regulated, tissue-dependent process rather than a purely mutation-driven event. By positioning interfacial mechanics between oncogenic mutant cells and their wild-type counterparts, alongside genetic and biochemical factors, our findings contribute to an emerging physical understanding of cancer, in which tissue mechanics play an active role in shaping oncogenic fate. Such a framework not only helps explain tissue-specific cancer susceptibility but also suggests that reinforcing mechanical barriers within epithelia may represent an underexplored strategy for limiting early tumor progression.

Removed in the revised version: Materials and methods

Removed in the revised version: Cell culture

Removed in the revised version: Non-transformed human mammary epithelial cell line MCF10A were maintained in complete medium composed of phenol-free Dulbecco’s modified Eagle’s medium (DMEM–F12, Gibco) supplemented with 5% charcoal-stripped horse serum (Gibco), 10 U ml−1 penicillin and 10 μg ml−1 streptomycin (Pen-Strep, Invitrogen), epidermal growth factor (20 ng/ml; PeproTech), hydrocortisone (0.5 mg/ml; Sigma-Aldrich), cholera toxin (100 ng/ml; Sigma-Aldrich), and insulin (10 μg/ml; Sigma-Aldrich) at 37°C in a humidified incubator with 5% CO2.

Removed in the revised version: Non-transformed human bronchial epithelium epithelial cell line BEAS2B were cultured in Dulbecco’s modified Eagle’s medium (DMEM, Gibco) supplemented with GlutaMax (Gibco) and 5% fetal bovine serum (FBS, Gibco) along with 10 U ml−1 penicillin and 10 μg ml−1 streptomycin (Pen-Strep, Invitrogen). Cells were maintained at 37 °C and 5% CO2 unless mentioned otherwise.

Removed in the revised version: Cell seeding

Removed in the revised version: 105 Cells were seeded on fibronectin-coated (10 μg ml-1) glass bottom dishes (35 mm Cellvis) and transfected once they reached 80% confluency (for singlets) and 60% confluency to get transfected groups.

Removed in the revised version: Transfection

Removed in the revised version: Cells were transiently transfected with the HrasV12-GFP expressing plasmid DNA (Addgene #18780) using Lipofectamine 3000 (Invitrogen), following the manufacturer’s protocol. Post-transfection, confluent monolayers were either fixed and immuno-stained or imaged directly.

Removed in the revised version: Blebbistatin addition

Removed in the revised version: Blebbistatin was diluted to 5 μM final solution and added 48h post-transfection. Samples were washed after 24 hours and fixed right away.

Removed in the revised version: Immunofluorescence

Removed in the revised version: Cells were fixed with 4% paraformaldehyde (ThermoFisher) diluted in 1X Phosphate buffered saline (PBS, Sigma-Aldrich) for 15 minutes at room temperature. After washing away the fixative with 1X PBS, cells were permeabilized with 0.3% TritonX-100 (Sigma-Aldrich) in 1X PBS. Nonspecific antibody binding was blocked by incubating the samples in blocking solution (2% Bovine Serum Albumin (BSA Sigma-Aldrich) and 2% FBS in PBS) at room temperature for 45 minutes. Further, cells were incubated with primary antibody prepared in blocking/staining solution overnight at 4°C. Post primary antibodies incubation, cells were washed thrice with 1X PBS for 5 minutes per wash and incubated with secondary antibodies conjugated to Alexa fluor 555, Alexa fluor 594 conjugated phalloidin and 4′,6-diamidino-2-phenylindole (DAPI) to mark the cell nucleus, prepared in blocking/staining solution, for 2 hours at room temperature. Finally, the samples were washed four times with 1X PBS for 5 minutes per wash before proceeding to microscopy.

Removed in the revised version: Microscopy

Removed in the revised version: Fluorescence images were acquired using 20X objective and 63x oil-immersion objective of a Zeiss Axio Observer 7 Inverted Microscope with a scientific-grade sCMOS camera by Iberoptics. For live imaging, samples were set up on a stage-top incubator and maintained at 37°C and 5% CO2 throughout imaging.

Removed in the revised version: Image Analysis

Removed in the revised version: Image analysis for this study was performed using Fiji [41] except bayesian force inference and particle image velocimetry (PIV) analysis, which was performed in MATLAB (MathWorks). Cellpose [42] was used to segment cells to get cell count, shape indices and actin intensitites for each cell while Tissue Analyzer [43] was used to segment cells for force inference.

Removed in the revised version: Quantification of cluster eccentricity and segregation index

Removed in the revised version: The term ‘eccentricity’ expresses a size reduction in the contact with neighboring cells. It was quantified using the formula e=(1−(ab)2, (where a= the length of the major axes and b= the length of the minor axes of the clusters), with a value closer to 0 indicating a more circular shape. Eccentricity quantifications were performed on fixed and live tissues using outlines that were manually drawn and measured using Fiji.

Removed in the revised version: The segregation index SI, defined as the average ratio of homotypic and all cell neighbors as done previously [44], quantifies the demixing degree. To quantify the SI, cells of each type were manually counted, and finally, to generate the plots, the SI was averaged over all GFP-labelled cells in one frame.

Removed in the revised version: Bayesian Force Inference

Removed in the revised version: Bayesian force inference was done as described previously [26] to determine the relative cellular pressures and cell-cell edge or line tension.

Removed in the revised version: Particle Image Velocimetry (PIV)

Removed in the revised version: Particle Image Velocimetry (PIV) was done using the PIV Lab package on MATLAB [45]. FFT window deformation (direct Fourier transform correlation with multiple passes and deforming windows) algorithm was used. Velocity maps were obtained by a 3-pass interrogation window of 128 × 128 pixels that halved after each pass. Tangential velocities were obtained by drawing concentric circles radially outward from the center to a point at the boundary of the mutant clusters and integrated over the entire circle using a custom python code. The output is the tangential velocity as a function of the circle.

Removed in the revised version: Statistical analysis

Removed in the revised version: Statistical analyses were carried out in GraphPad Prism 10. Statistical significance was calculated by Unpaired t-test with Welch’s correction. p-values greater than 0.05 were considered to be statistically not significant. Each plot has data pooled from multiple frames of one out of three independent experiments. Violin plots have median shown as a bold dashed line and the first and third quartiles shown as thin dashed lines. Box and whisker plots have whiskers extending from the minimum to the maximum values, showing the full range of the dataset including outliers. Scatter-bar plots were displayed as mean ± s.e.m. No statistical methods were used to set the sample size. Quantification was done using data from at least three independent biological replicates. For analysis involving live-imaging experiments, data were collected from three independent experiments. All the experiments with representative images were repeated at least three times.

Removed in the revised version: Vertex Model

Removed in the revised version: A bi-disperse model is used to demonstrate the role of interfacial tension between two types of cells, specifically the HRasV12 mutant and wild-type cells. EVM=∑α∈WT[KαWT2(Aα−Aα0,WT)2+ΓαWT2(Pα−Pα0,WT)2]+∑α∈M[KαM2(Aα−Aα0,M)2+∑αΓαM2(Pα−Pα0,M)2]+∑<ij>Λijlij The first two sums are included in the standard vertex model, whereas the last term is introduced in our model to account for interfacial interactions between cells. Specifically, Λij = 0 for interfaces between same cell types, and Λij = Λ between different cell types. Using a line tension term like Λlij in the vertex model energy, where the lij is the length of an interface edge, we can obtain both segregation and spreading dynamics just from energy minimization. This term has been described as heterotypic line tension (HLT)[46]. Due to energy minimization, heterotypic interface length is minimized for Λ > 0 and maximized for Λ < 0, leading to segregation of the two cell types in the former case and complete mixing of them in the latter case.

Removed in the revised version: The mutant cluster in the experiments with MCF10A cells showed an actin belt forming along the cluster interface, with the cells along the interface being more elongated than normal. This actin ring is important for the compaction of mutant clusters in the presence of wild-type cells in the epithelial monolayer. An isotropic interfacial tension, such as the Λlij term in our model cannot capture such anisotropic behaviour. Recently, it has been proposed that such anisotropic behaviour can arise from the anisotropic distribution of the cytoskeletal filaments inside the cell[40]. The resulting active interfacial tension, called the shape-tension coupling, can generate elongated shapes as observed in the experiments. The elongated cells in the experiments indeed have anisotropic distributions of the stress fibres. Hence, we included the effect of shape-tension in our model. The total interfacial tension is Λij + γij(t; Π0), where Λij is a time-independent, passive, isotropic tension and γij is a time-dependent active tension that depends on the shape of the cells sharing the interface and requires a threshold stress Π0 to be nonzero. The details of the shape-tension coupling is described in the interface. We find that a nonzero activation stress Π0 is necessary to reproduce experimentally observed results.